Before you turn this assignment in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says "YOUR ANSWER HERE" or `YOUR CODE HERE` and remove the `raise NotImplementedError()` lines. 

Code blocks starting with a `# tests` comment provide unit tests which have to run without errors in order to get full points. Be aware that there might be further 'secret' tests to check correct implementation! I.e. the provided unit tests are necessary but not sufficient for full points!

You are always welcome to add **additional plots, tests, or debug outputs**.
However, make sure to: **1) not break the automated tests**, and **2) switch off any excessive debug output** when you submit your notebook!

Please add your name and student ID below:

<!-- grade_id: cell-bc6886da45bcfc76, points: 0, grade: false, solution: true -->
```python
NAME = "Konstantinos Vakalopoulos" 
STUDENT_ID = "12223236" 
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: name_and_id_test, points: 0, grade: true, solution: false -->
```python
assert len(NAME) > 0, "Enter your name!"
assert len(STUDENT_ID) > 0, "Enter your student ID!"
```


# Intelligent Audio and Music Analysis Assignment 0

The main purpose of this assignment is to get you familiarized with the infrastructure and workflow.  

## Task 1: Import numpy

Load the python math library [numpy](https://numpy.org/doc/stable/) into namespace variable `np`.

<!-- grade_id: import_numpy, points: 0, grade: false, solution: true -->
```python
# load the python math library numpy into a local namespace variable called "np"
# this is achieved using the import <lib> as <name> syntax
# YOUR CODE HERE
import numpy as np
```

<!-- grade_id: import_numpy_test, points: 1, grade: true, solution: false -->
```python
# test if numpy has been loaded successfully

np.version.version
print('All tests successful!')
```

Output:
```
All tests successful!
```

## Task 2: Create a Matrix

Create a matrix ([ndarray](https://numpy.org/doc/stable/reference/generated/numpy.ndarray.html)) `x` of random numbers 
with size 4x4.

<!-- grade_id: rand_x, points: 0, grade: false, solution: true -->
```python
# numpy has a module called "random", use the rand function. 
# See documentation at: 
# https://numpy.org/doc/stable/reference/random/generated/numpy.random.rand.html

# YOUR CODE HERE
from numpy import random
x = random.rand(4,4)

print("x = {}".format(x))
```

Output:
```
x = [[0.60401524 0.51029294 0.74465668 0.53505097]
 [0.26602286 0.64578061 0.31781944 0.65441859]
 [0.74613021 0.78059703 0.69286132 0.5116467 ]
 [0.74767905 0.87145249 0.221681   0.29565113]]
```

<!-- grade_id: rand_x_test, points: 1, grade: true, solution: false -->
```python
# test x for properties:

assert x.shape == (4, 4)
assert all([0 <= ex <= 1 for ex in x.flatten()])
print('All tests successful!')
```

Output:
```
All tests successful!
```

<!-- grade_id: rand_x_test_exc, points: 1, grade: true, solution: false -->
```python
# test for exception: 

try:
    x[10, 10]
except IndexError:
    pass
else:       
    raise AssertionError('did not raise ValueError')
print('All tests successful!')
```

Output:
```
All tests successful!
```

## Task 3: Elementwise Matrix Operations
Calculate the *elementwise* square of `x` into a new variable named `y`.

<!-- grade_id: x_mult, points: 0, grade: false, solution: true -->
```python
# Arrays can be used as matrices, thus we can choose between matrix multiplication and elementwise multiplication.
# https://numpy.org/doc/stable/reference/generated/numpy.matmul.html
# https://numpy.org/doc/stable/reference/generated/numpy.multiply.html

# YOUR CODE HERE
y = np.square(x)
```

<!-- grade_id: x_mult_test, points: 1, grade: true, solution: false -->
```python
# Test shape of y
# Note: for full points, values of y must also be correct.

assert y.shape == (4, 4)

print('All tests successful!')
```

Output:
```
All tests successful!
```

## Congratulations, you are done!

Reminder:
Before you turn this assignment in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).